---
title: 'Lesson 2.1: Intro to normalization'
description: 'Chapter description goes here.'
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Lesson 2.1: Intro to normalization

```yaml
type: VideoExercise
key: 325150b3d9
xp: 50
```

`@projector_key`
35a9424dc6a263187b08bcf4cc18ccab

---

## Normalizing sales data

```yaml
type: PureMultipleChoiceExercise
key: 270526a1c9
xp: 50
```



`@hint`


`@possible_answers`


`@feedback`
